# Sprint Challenge
## *Data Science Unit 4 Sprint 1*

After a week of Natural Language Processing, you've learned some cool new stuff: how to process text, how turn text into vectors, and how to model topics from documents. Apply your newly acquired skills to one of the most famous NLP datasets out there: [Yelp](https://www.yelp.com/dataset). As part of the job selection process, some of my friends have been asked to create analysis of this dataset, so I want to empower you to have a head start.  

The real dataset is massive (almost 8 gigs uncompressed). I've sampled the data for you to something more manageable for the Sprint Challenge. You can analyze the full dataset as a stretch goal or after the sprint challenge. As you work on the challenge, I suggest adding notes about your findings and things you want to analyze in the future.

## Challenge Objectives
Successfully complete all these objectives to earn full credit. 

**Successful completion is defined as passing all the unit tests in each objective.**  

Each unit test that you pass is 1 point. 

There are 5 total possible points in this sprint challenge. 


There are more details on each objective further down in the notebook.*
* <a href="#p1">Part 1</a>: Write a function to tokenize the yelp reviews
* <a href="#p2">Part 2</a>: Create a vector representation of those tokens
* <a href="#p3">Part 3</a>: Use your tokens in a classification model on yelp rating
* <a href="#p4">Part 4</a>: Estimate & Interpret a topic model of the Yelp reviews

____

# Before you submit your notebook you must first

1) Restart your notebook's Kernel

2) Run all cells sequentially, from top to bottom, so that cell numbers are sequential numbers (i.e. 1,2,3,4,5...)
- Easiest way to do this is to click on the **Cell** tab at the top of your notebook and select **Run All** from the drop down menu. 

3) Comment out the cell that generates a pyLDAvis visual in objective 4 (see instructions in that section). 
____



### Import Data

<!-- grade_id: cell-395851cd95d17235, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# Load reviews from URL
data_url = 'https://raw.githubusercontent.com/LambdaSchool/data-science-practice-datasets/main/unit_4/unit1_nlp/review_sample.json'

# Import data into a DataFrame named df
# YOUR CODE HERE
df = pd.read_json(data_url, lines=True)
```

<!-- grade_id: cell-cb5006475e42b8f9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert isinstance(df, pd.DataFrame), 'df is not a DataFrame. Did you import the data into df?'
assert df.shape[0] == 10000, 'DataFrame df has the wrong number of rows.'
```

```python
df.head()
```

Output:
```
business_id  cool                date  funny  \
0  nDuEqIyRc8YKS1q1fX0CZg     1 2015-03-31 16:50:30      0   
1  eMYeEapscbKNqUDCx705hg     0 2015-12-16 05:31:03      0   
2  6Q7-wkCPc1KF75jZLOTcMw     1 2010-06-20 19:14:48      1   
3  k3zrItO4l9hwfLRwHBDc9w     3 2010-07-13 00:33:45      4   
4  6hpfRwGlOzbNv7k5eP9rsQ     1 2018-06-30 02:30:01      0   

                review_id  stars  \
0  eZs2tpEJtXPwawvHnHZIgQ      1   
1  DoQDWJsNbU0KL1O29l_Xug      4   
2  DDOdGU7zh56yQHmUnL1idQ      3   
3  LfTMUWnfGFMOfOIyJcwLVA      1   
4  zJSUdI7bJ8PNJAg4lnl_Gg      4   

                                                text  useful  \
0  BEWARE!!! FAKE, FAKE, FAKE....We also own a sm...      10   
1  Came here for lunch Togo. Service was quick. S...       0   
2  I've been to Vegas dozens of times and had nev...       2   
3  We went here on a night where they closed off ...       5   
4  3.5 to 4 stars\n\nNot bad for the price, $12.9...       5   

                  user_id  
0  n1LM36qN
```

## Part 1: Tokenize Function
<a id="#p1"></a>

Complete the function `tokenize`. Your function should
- accept one document at a time
- return a list of tokens

You are free to use any method you have learned this week.

```python
# Optional: Consider using spaCy in your function. The spaCy library can be imported by running this cell.
# A pre-trained model (en_core_web_sm) has been made available to you in the CodeGrade container.
# If you DON'T need use the en_core_web_sm model, you can comment it out below.
# import spacy
# nlp = spacy.load('en_core_web_sm')
```

<!-- grade_id: cell-3d570d5a1cd6cb64, points: 0, grade: false, solution: true -->
```python
import re

def tokenize(doc):
    doc = doc.replace('\n', ' ')
    doc = re.sub('[^a-zA-Z ]', '', doc)
    tokens = doc.strip().lower()
    tokens = [token for token in tokens.split() if len(token) > 2]

    return tokens
```

<!-- grade_id: cell-02da164f6fbe730a, points: 1, grade: true, solution: false -->
```python
'''Testing'''
assert isinstance(tokenize(df.sample(n=1)["text"].iloc[0]), list), "Make sure your tokenizer function accepts a single document and returns a list of tokens!"
```

## Part 2: Vector Representation
<a id="#p2"></a>
1. Create a vector representation of the reviews (i.e. create a doc-term matrix).
2. Write a fake review and query for the 10 most similar reviews, print the text of the reviews. Do you notice any patterns?
    - Given the size of the dataset, use `NearestNeighbors` model for this. 

```python
from sklearn.feature_extraction.text import CountVectorizer, TfidfVectorizer
```

<!-- grade_id: cell-0e96491cb529202c, points: 0, grade: false, solution: true -->
```python
%%time
# Create a vector representation of the reviews 
# Name that doc-term matrix "dtm"

# YOUR CODE HERE
tfidf = TfidfVectorizer(
    stop_words='english',
    max_features=5000,
    tokenizer=tokenize
)
dtm = tfidf.fit_transform(df['text'])
dtm = pd.DataFrame(dtm.todense(), columns=tfidf.get_feature_names())
```

Output:
```
CPU times: user 774 ms, sys: 127 ms, total: 901 ms
Wall time: 998 ms
```

```python
dtm.head()
```

Output:
```
ability  able  absolute  absolutely  accent  accept  acceptable  accepted  \
0      0.0   0.0       0.0         0.0     0.0     0.0         0.0       0.0   
1      0.0   0.0       0.0         0.0     0.0     0.0         0.0       0.0   
2      0.0   0.0       0.0         0.0     0.0     0.0         0.0       0.0   
3      0.0   0.0       0.0         0.0     0.0     0.0         0.0       0.0   
4      0.0   0.0       0.0         0.0     0.0     0.0         0.0       0.0   

   access  accessible  ...  yuck  yuk  yum  yummy  yup  zen  zero  zone  zoo  \
0     0.0         0.0  ...   0.0  0.0  0.0    0.0  0.0  0.0   0.0   0.0  0.0   
1     0.0         0.0  ...   0.0  0.0  0.0    0.0  0.0  0.0   0.0   0.0  0.0   
2     0.0         0.0  ...   0.0  0.0  0.0    0.0  0.0  0.0   0.0   0.0  0.0   
3     0.0         0.0  ...   0.0  0.0  0.0    0.0  0.0  0.0   0.0   0.0  0.0   
4     0.0         0.0  ...   0.0  0.0  0.0    0.0  0.0  0.0   0.0   0.0  0.0   

   zucchini  
0       0.0  
1       0.0  
```

<!-- grade_id: cell-3d5bc610a8ec6b24, points: 0, grade: false, solution: true -->
```python
# Create and fit a NearestNeighbors model named "nn"
from sklearn.neighbors import NearestNeighbors

# YOUR CODE HERE
nn = NearestNeighbors(n_neighbors=10, algorithm='kd_tree')
nn.fit(dtm)
```

Output:
```
NearestNeighbors(algorithm='kd_tree', n_neighbors=10)
```

<!-- grade_id: cell-c43704dcff67e99b, points: 1, grade: true, solution: false -->
```python
'''Testing.'''
assert nn.__module__ == 'sklearn.neighbors._unsupervised', ' nn is not a NearestNeighbors instance.'
assert nn.n_neighbors == 10, 'nn has the wrong value for n_neighbors'
```

<!-- grade_id: cell-496203e8746296ca, points: 0, grade: false, solution: true -->
```python
# Create a fake review and find the 10 most similar reviews

# YOUR CODE HERE
positive = ['I loved this place so much that I decided I am going to give birth to all my kids here.']
negative = ['I hated this place so much that I decided to commit vehicular manslaughter.']
pinput = tfidf.transform(positive).todense()
ninput = tfidf.transform(negative).todense()
poutput = nn.kneighbors(pinput, n_neighbors=10, return_distance=False)
noutput = nn.kneighbors(ninput, n_neighbors=10, return_distance=False)
```

```python
for x in poutput[0]:
    print(df['text'][x], '\n')
```

Output:
```
天氣很熱吃不下東西，今天我點了一個韓國冷面湯、餐後點了甜點，冰沙系列不會太甜膩，覺得店家很用心製作，包含擺盤精緻、佐料衛生，夏日想開胃，這是一個不錯的選擇，服務人員也很敬業，以後會常常來 

O  o  thenk 6nnn  .b  cgv  xx TV cvg  9 nvehxcfvvv3c nb b  c  y4  nb and the vghvhridd h 0d  c       v   3,  vv  4  ruddy 

旅行でラスベガスに来ましたがネイルがはげてるのが気になり、探したお店でした。
質問にも丁寧に答えてくれましたし、日本人の方も日本語が話せる方も居て、とても綺麗で居心地のいいお店でした。 
ネイルはちはるさんと言う綺麗な方が丁寧にしてくれとても気に入りました。
予定になかったまつ毛エクステもお願いし、日本ではまだあまりないブラウンカラーのエクステをしてもらい、とても気に入りました。
また是非マッサージなどで伺いたいと思います。 

By far the best public pool to take your kids. The 2 giant slides (for 42" kids), a small lazy river and a water vortex which is a cool and unique feature that my kids could not get enough of. For $1 per kid, I cannot think of a better way to spend a day. It is well run and organized and the teenage staff are all delightful, friendly, clean cut kids which I loved as good role modeling for my grade school kids. Lots of seating to have lunch which we brought in. 

I have had both my kids in this school. My daughter is five and my son is 11. My son now has a s
```

```python
for x in noutput[0]:
    print(df['text'][x], '\n')
```

Output:
```
旅行でラスベガスに来ましたがネイルがはげてるのが気になり、探したお店でした。
質問にも丁寧に答えてくれましたし、日本人の方も日本語が話せる方も居て、とても綺麗で居心地のいいお店でした。 
ネイルはちはるさんと言う綺麗な方が丁寧にしてくれとても気に入りました。
予定になかったまつ毛エクステもお願いし、日本ではまだあまりないブラウンカラーのエクステをしてもらい、とても気に入りました。
また是非マッサージなどで伺いたいと思います。 

O  o  thenk 6nnn  .b  cgv  xx TV cvg  9 nvehxcfvvv3c nb b  c  y4  nb and the vghvhridd h 0d  c       v   3,  vv  4  ruddy 

天氣很熱吃不下東西，今天我點了一個韓國冷面湯、餐後點了甜點，冰沙系列不會太甜膩，覺得店家很用心製作，包含擺盤精緻、佐料衛生，夏日想開胃，這是一個不錯的選擇，服務人員也很敬業，以後會常常來 

Man, I hated the ambiance. The outdoor seating looked nice, but indoors it's tacky and plain. Service and the food were both average. Would absolutely not recommend this place. 

I was looking for a new massage/facial place after moving from Chicago. I tried one and hated it. After trying Mozaik I knew I'd be back. I love this place!!! Clean, comfortable, and chic. Plus great service and friendly people. 

I decided to try this place. Everyone was friendly however I would never go back.  I got the worst fungus in my fingernail. 

I stayed here for 5 nights 
```

## Part 3: Classification
<a id="#p3"></a>
Your goal in this section will be to predict `stars` from the review dataset. 

1. Create a pipeline object with a sklearn `CountVectorizer` or `TfidfVector` and any sklearn classifier.
    - Use that pipeline to train a model to predict the `stars` feature (i.e. the labels). 
    - Use that Pipeline to predict a star rating for your fake review from Part 2. 



2. Create a parameter dict including `one parameter for the vectorizer` and `one parameter for the model`. 
    - Include 2 possible values for each parameter
    - **Use `n_jobs` = 1** 
    - Due to limited computational resources on CodeGrader `DO NOT INCLUDE ADDITIONAL PARAMETERS OR VALUES PLEASE.`
    
    
3. Train the entire pipeline with a GridSearch
    - Name your GridSearch object as `gs`

```python
X = df['text']
y = df['stars']
```

<!-- grade_id: cell-e2beb0252d274bba, points: 0, grade: false, solution: true -->
```python
from sklearn.pipeline import Pipeline
# from sklearn.feature_extraction.text import TfidfVectorizer
from sklearn.neighbors import KNeighborsClassifier
from sklearn.model_selection import GridSearchCV

# Name the gridsearch instance "gs"

# YOUR CODE HERE
clf = KNeighborsClassifier()
pipe = Pipeline([('vect', tfidf), ('clf', clf)])

params = {
    'vect__max_df': [.75, .85],
    'clf__n_neighbors': [7, 15]
}

gs = GridSearchCV(pipe, params, n_jobs=1, cv=3, verbose=1)
gs.fit(X, y)
```

Output:
```
Fitting 3 folds for each of 4 candidates, totalling 12 fits
[Parallel(n_jobs=1)]: Using backend SequentialBackend with 1 concurrent workers.
[Parallel(n_jobs=1)]: Done  12 out of  12 | elapsed:   23.0s finished
GridSearchCV(cv=3,
             estimator=Pipeline(steps=[('vect',
                                        TfidfVectorizer(max_features=5000,
                                                        stop_words='english',
                                                        tokenizer=<function tokenize at 0x7fd10bbbedc0>)),
                                       ('clf', KNeighborsClassifier())]),
             n_jobs=1,
             param_grid={'clf__n_neighbors': [7, 15],
                         'vect__max_df': [0.75, 0.85]},
             verbose=1)
```

```python
print(gs.best_score_)
gs.best_params_
```

Output:
```
0.513701010173003
{'clf__n_neighbors': 15, 'vect__max_df': 0.75}
```

<!-- grade_id: cell-d07134c6fe5d056e, points: 1, grade: true, solution: false -->
```python
# Visible Testing
prediction = gs.predict(["I wish dogs knew how to speak English."])[0]
assert prediction in df.stars.values, 'You gs object should be able to accept raw text within a list. Did you include a vectorizer in your pipeline?'
```

## Part 4: Topic Modeling

Let's find out what those yelp reviews are saying! :D

1. Estimate a LDA topic model of the review text
    - Set num_topics to `5`
    - Name your LDA model `lda`
2. Create 1-2 visualizations of the results
    - You can use the most important 3 words of a topic in relevant visualizations. Refer to yesterday's notebook to extract. 
3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

When you instantiate your LDA model, it should look like this: 

```python
lda = LdaModel(corpus=corpus,
               id2word=id2word,
               random_state=723812,
               num_topics = num_topics,
               passes=1
              )

```

__*Note*__: You can pass the DataFrame column of text reviews to gensim. You do not have to use a generator.

## Note about  pyLDAvis

**pyLDAvis** is the Topic modeling package that we used in class to visualize the topics that LDA generates for us.

You are welcomed to use pyLDAvis if you'd like for your visualization. However, **you MUST comment out the code that imports the package and the cell that generates the visualization before you submit your notebook to CodeGrade.** 

Although you should leave the print out of the visualization for graders to see (i.e. comment out the cell after you run it to create the viz). 

```python
from gensim import corpora
# Due to limited computationalresources on CodeGrader, use the non-multicore version of LDA 
from gensim.models.ldamodel import LdaModel
import gensim
import re
```

### 1. Estimate a LDA topic model of the review tex

<!-- grade_id: cell-66331a185ff52f15, points: 0, grade: false, solution: true -->
```python
# Remember to read the LDA docs for more information on the various class attirbutes and methods available to you
# in the LDA model: https://radimrehurek.com/gensim/models/ldamodel.html

# don't change this value 
num_topics = 5

# use tokenize function you created earlier to create tokens
tokenized = [tokenize(text) for text in df['text']]

# create a id2word object (hint: use corpora.Dictionary)
id2word = corpora.Dictionary(tokenized)

# create a corpus object (hint: id2word.doc2bow)
corpus = [id2word.doc2bow(text) for text in tokenized]

# instantiate an lda model
lda = LdaModel(
    corpus=corpus,
    id2word=id2word,
    random_state=42,
    num_topics=num_topics,
    passes=1
)
```

#### Testing

<!-- grade_id: cell-5a3c181311134fa9, points: 1, grade: true, solution: false -->
```python
# Visible Testing
assert lda.get_topics().shape[0] == 5, 'Did your model complete its training? Did you set num_topics to 5?'
```

#### 2. Create 1-2 visualizations of the results

<!-- grade_id: cell-9b043e992fbd218c, points: 0, grade: false, solution: true -->
```python
# import seaborn as sns
# import matplotlib.pyplot as plt
# import pyLDAvis
# import pyLDAvis.gensim

# # Use pyLDAvis (or a ploting tool of your choice) to visualize your results 

# # YOUR CODE HERE
# pyLDAvis.enable_notebook()
# vis = pyLDAvis.gensim.prepare(lda, corpus, id2word)
# vis
```

#### 3. In markdown, write 1-2 paragraphs of analysis on the results of your topic model

Out of the 5 topics that were generated, we obviously have our positive and negative feedback topics. These contained postive and negative words like 'love' and 'hate' and similar, no surprise there. A topic that was interesting is that it managed to pick out reviews where the topic was children. Contained words like 'kid' and 'education' and the such. There was a topic that seemed to be focused on food, which as a food lover myself I can personally relate to. The last topic was seemingly focused on employees or staff, maybe the service.